<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Assignment-#4:-Feature-Engineering" data-toc-modified-id="Assignment-#4:-Feature-Engineering-1">Assignment #4: Feature Engineering</a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-2">Learning Outcomes</a></span></li><li><span><a href="#Data-Science-Workflow:" data-toc-modified-id="Data-Science-Workflow:-3">Data Science Workflow:</a></span></li><li><span><a href="#1)-Ask:-" data-toc-modified-id="1)-Ask:--4">1) Ask: </a></span></li><li><span><a href="#2)-Acquire" data-toc-modified-id="2)-Acquire-5">2) Acquire</a></span></li><li><span><a href="#3)-Process" data-toc-modified-id="3)-Process-6">3) Process</a></span></li><li><span><a href="#4)-Model" data-toc-modified-id="4)-Model-7">4) Model</a></span></li><li><span><a href="#5)-Deliver" data-toc-modified-id="5)-Deliver-8">5) Deliver</a></span></li><li><span><a href="#Hints-&amp;-Guidelines" data-toc-modified-id="Hints-&amp;-Guidelines-9">Hints &amp; Guidelines</a></span></li></ul></div>

Assignment #4: Feature Engineering
----
<br>
<br>
<center><img src="https://vangogh.teespring.com/og_pic/78661465/15927516/front.jpg" width="75%"/></center>

Learning Outcomes
------

__By the end of this session, you should be able to__:


- Define a Pipeline to automate the machine learning process for processing data, including imputation and encoding.
- Fit a Random Forest™ Regressor to predict real-world data.


Data Science Workflow:
----
1. Ask
2. Acquire
3. Process
4. Model
5. Deliver 

1) Ask: 
----

Can we predict the duration of bike trips for [Bay Wheels](https://www.lyft.com/bikes/bay-wheels)?

2) Acquire
----

The data comes from https://www.lyft.com/bikes/bay-wheels/system-data.

Each trip is anonymized and features used are:

- Trip ID
- Start Station ID
- Start Station Name
- End Station ID
- End Station Name
- Bike ID
- User Type (Subscriber or Customer – “Subscriber” = Member or “Customer” = Casual)
- Member Year of Birth
- Member Gender

The targets is Trip Duration (in seconds).

3) Process
-----

```python
reset -fs
```

```python
import pandas as pd
```

```python
path = "../../data/random_forest_regression/"

# Load data
X_train = pd.read_csv(path+"X_train.csv")
X_test= pd.read_csv(path+"X_test.csv")
y_train = pd.read_csv(path+"y_train.csv", header=None)
y_train = y_train.values.ravel()
```

Define a pipeline to encapsulate the entire transformation and estimator steps.

```python
from sklearn.pipeline import Pipeline
```

You'll need one process for numeric features and another process for categorical features using [Column Transformer with Mixed Types](https://scikit-learn.org/stable/auto_examples/compose/plot_column_transformer_mixed_types.html).

```python
from sklearn import compose
```

Since this is real-world data, there are missing data. You'll have to impute.

```python
from sklearn.experimental import enable_iterative_imputer
from sklearn import impute
```

All transformations must be done in scikit-learn.

```python
from sklearn import preprocessing
```

__No pandas__

```python
# 👋 Pandas 
del(pd)
```

The imports are at the module level so you'll have write code in this way `module.Class()`

Examples:

```python
compose.ColumnTransformer()
```

or

```python
impute.SimpleImputer()
```

or

```python
preprocessing.OneHotEncoder()
```

<!-- grade_id: cell-a18ca1a5922c26e5, points: 0, grade: false, solution: true -->
```python
def make_pipeline(regressor=None):
    "Create a single pipeline that processing the data and then fits the regressor." 
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return pipeline
```

4) Model
-----

Fit only RandomForestRegressor __without__ hyperparameter tuning.

```python
from sklearn.ensemble import RandomForestRegressor

regressor = RandomForestRegressor(n_estimators=30,
                                  min_samples_leaf=1,
                                  random_state=42)
```

<!-- grade_id: cell-dd711355aa43bf7d, points: 3, grade: true, solution: false -->
```python
"""
3 points
Test code for the 'make_pipeline' function. 
This cell should NOT give any errors when it is run.
"""

pipeline = make_pipeline(regressor)

assert type(pipeline) == Pipeline
assert len(pipeline) == 2
```

```python
# Train Estimator
pipeline.fit(X_train, y_train);
```

Let's use `median_absolute_error` as our North Star metric.

```python
from sklearn.metrics import median_absolute_error
```

```python
y_pred = pipeline.predict(X_train)
medae_value_train = median_absolute_error(y_train, y_pred)
print(f"{medae_value_train:.4f} medae on train dataset")  
```

Let's use [cross_val_score](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.cross_val_score.html) to get a sense of the Variance of our system (i.e., features, methodology, algorithm, and hyperparameters.)

```python
from sklearn.metrics import make_scorer
from sklearn.model_selection import cross_val_score
```

```python
median_absolute_error_scorer = make_scorer(median_absolute_error)
cross_val_score(pipeline, 
                X_train, 
                y_train, 
                scoring=median_absolute_error_scorer,
                cv=5)  
```

5) Deliver
-----

How well can your model predict the test data?

```python
# Check that it works for the test data
y_pred = pipeline.predict(X_test)
```

<!-- grade_id: cell-e26b2982bc66dc84, points: 1, grade: true, solution: false -->
```python
"""
1 point for median_absolute_error lower than 90
Tests are hidden since you don't have access to test targets.
"""

```

<!-- grade_id: cell-70a214c26f5dcbc6, points: 1, grade: true, solution: false -->
```python
"""
1 point for median_absolute_error lower than 85
Tests are hidden since you don't have access to test targets.
"""

```

<!-- grade_id: cell-78ae6f811561e0c3, points: 1, grade: true, solution: false -->
```python
"""
1 point for median_absolute_error lower than 80
Tests are hidden since you don't have access to test targets.
"""

```

<!-- grade_id: cell-c55fe75be033dd39, points: 1, grade: true, solution: false -->
```python
"""
1 point for median_absolute_error lower than 75
Tests are hidden since you don't have access to test targets.
"""

```

<!-- grade_id: cell-c4b5557adf91c401, points: 1, grade: true, solution: false -->
```python
"""
1 point for median_absolute_error lower than 70
Tests are hidden since you don't have access to test targets.
"""

```

<!-- grade_id: cell-7b2c53fa571837d2, points: 1, grade: true, solution: false -->
```python
"""
1 point for median_absolute_error lower than 65
Tests are hidden since you don't have access to test targets.
"""

```

<!-- grade_id: cell-75d173a05606c6e3, points: 1, grade: true, solution: false -->
```python
"""
1 point for median_absolute_error lower than 60
Tests are hidden since you don't have access to test targets.
"""

```

Hints & Guidelines
-----

- No other imports allowed. 
- You get access to all of `impute`, `compose`, and `preprocessing`
- Think about the best way to encode different types of features.