<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Assignment-#3:-Multiclass-Classification,-Pipelines,-and-Hyperparameter-Tuning" data-toc-modified-id="Assignment-#3:-Multiclass-Classification,-Pipelines,-and-Hyperparameter-Tuning-1">Assignment #3: Multiclass Classification, Pipelines, and Hyperparameter Tuning</a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-2">Learning Outcomes</a></span></li><li><span><a href="#Data-Science-Workflow:" data-toc-modified-id="Data-Science-Workflow:-3">Data Science Workflow:</a></span></li><li><span><a href="#1)-Ask:-" data-toc-modified-id="1)-Ask:--4">1) Ask: </a></span></li><li><span><a href="#2)-Acquire" data-toc-modified-id="2)-Acquire-5">2) Acquire</a></span></li><li><span><a href="#3)-Process" data-toc-modified-id="3)-Process-6">3) Process</a></span></li><li><span><a href="#Load-Data" data-toc-modified-id="Load-Data-7">Load Data</a></span></li><li><span><a href="#4)-Model" data-toc-modified-id="4)-Model-8">4) Model</a></span></li><li><span><a href="#5)-Deliver" data-toc-modified-id="5)-Deliver-9">5) Deliver</a></span></li><li><span><a href="#Random-Search" data-toc-modified-id="Random-Search-10">Random Search</a></span></li><li><span><a href="#Hints:" data-toc-modified-id="Hints:-11">Hints:</a></span></li></ul></div>

Assignment #3: Multiclass Classification, Pipelines, and Hyperparameter Tuning
----
<br>
<br>
<center><img src="https://imgs.xkcd.com/comics/data_pipeline.png" width="75%"/></center>

> "Is the pipeline literally running from your laptop?"   
> "Don't be silly, my laptop disconnects far too often to host a service we rely on. It's running on my phone."


Learning Outcomes
------

__By the end of this session, you should be able to__:


- Define a Pipeline to automate the machine learning process for multiple Estimators.
- Perform Random Search on an Estimator to improve performance.


Data Science Workflow:
----
1. Ask
2. Acquire
3. Process
4. Model
5. Deliver 

1) Ask: 
----

Which classification algorithm can learn to best predict a mulitclass targets in a tabular dataset?

In particular, predict the type of glass based on physical characteristics.

2) Acquire
----

The classification of glass types is motivated by crime scene investigation, often broken glass is left behind after a crime. The ability to identify the type of glass can help elucidate what happened at a crime scene.

The features are:

- RI: refractive index
- Na: Sodium 
- Mg: Magnesium
- Al: Aluminum
- Si: Silicon
- K: Potassium
- Ca: Calcium
- Ba: Barium
- Fe: Iron

The targets are 6 discrete types of glass.

3) Process
-----

```python
reset -fs
```

```python
import numpy as np
import pandas as pd
```

Load Data
----



```python
path = "../../data/multiclass_classification/"

# Load train data
X_train = pd.read_csv(path+"X_train.csv")
y_train = pd.read_csv(path+"y_train.csv")
y_train = y_train.values.ravel()
```

4) Model
-----

Fit the following models. No other models are allowed.

```python
from sklearn.linear_model import LogisticRegression
from sklearn.neighbors import KNeighborsClassifier
from sklearn.naive_bayes import GaussianNB
from sklearn.svm import SVC
from sklearn.ensemble import RandomForestClassifier
```

```python
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import StandardScaler
```

<!-- grade_id: cell-6f6346a5be240c31, points: 0, grade: false, solution: true -->
```python
def make_pipelines():
    """Create a pipeline for each of the following algorithms:
    1. Logistic Regression
    2. k-nearest neighbors (KNN) 
    3. Naive Bayes (Guassian)
    4. Support Vector Machines (SVM)
    5. Random Forest™ 
    
    If appropriate, apply StandardScaler before the algorithm.   
    Use default hyperparameters.
    If an algorithm takes random_state then random_state=42 
    
    Return a list of all the pipelines.
    """ 
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return pipelines
```

<!-- grade_id: cell-40fd8590b6d782a1, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the 'make_pipelines' function. 
This cell should NOT give any errors when it is run.
"""

pipelines = make_pipelines()

assert type(pipelines) == list
assert len(pipelines) == 5
assert all(type(pipe) == Pipeline for pipe in pipelines)
```

<!-- grade_id: cell-cf0b582e0a12220a, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test if StandardScaler is applied before each algorithms that requires it.
Hidden tests.
"""

```

```python
# Train all the models
for pipe in pipelines:
    pipe.fit(X_train, y_train)
```

<!-- grade_id: cell-a5d2f54024afc4a9, points: 0, grade: false, solution: true -->
```python
# Write code to look at confusion matrix for each model (ungraded)

from sklearn.metrics import confusion_matrix

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4b1ed1f9cffd6bd2, points: 0, grade: false, solution: true -->
```python
# Write code to look at classification_report for each model (ungraded)

from sklearn.metrics import classification_report

# YOUR CODE HERE
raise NotImplementedError()
```

Weighted F1 score will be our Northstar metric.

Sort each fitted model by weighted F1 score.

<!-- grade_id: cell-60229cf2dac50958, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import f1_score

def sort_models(pipelines, X_data, y_data, metric=f1_score, average='weighted'):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return scores_sorted
```

<!-- grade_id: cell-e2a4c7e719826aad, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the 'sort_models' function. 
This cell should NOT give any errors when it is run.
"""

scores_sorted_train = sort_models(pipelines, X_train, y_train, metric=f1_score, average='weighted')

assert type(scores_sorted_train) == dict
assert set(scores_sorted_train.keys()) == {'GaussianNB', 'KNeighborsClassifier', 'LogisticRegression', 'RandomForestClassifier', 'SVC'}
assert all(type(value) == np.float64 for value in scores_sorted_train.values())
```

<!-- grade_id: cell-50573d2f47fc949c, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test if models are sorted from best-to-worst on training data.
Hidden tests.
"""

# Example: GaussianNB will be the worst model.
assert list(scores_sorted_train.keys())[-1] == 'GaussianNB'

```

5) Deliver
-----

How well can your model predict the test data?

```python
# Load test data
X_test= pd.read_csv(path+"X_test.csv")
```

<!-- grade_id: cell-bbf574ba9f26fdb0, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test if models are sorted from best-to-worst on test data.
Hidden tests.
"""

```

Random Search
-----

As discussed in class, Random Search is much faster than Grid Search. Random Search can also yield similar performance on evaluation metrics.

Take your best algorithm from above and tune it with random search.

```python
from sklearn.model_selection import RandomizedSearchCV

# Limit the number of folds and samples

cv=5
n_iter=20
```

<!-- grade_id: cell-b9d686fff15f5902, points: 0, grade: false, solution: true -->
```python
def make_random_cv():
    """
    Steps:
    1. Define Estimator
    2. Define hyperparameter search space
    3. Instantiate RandomizedSearchCV
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return clf_random_cv
```

<!-- grade_id: cell-f2fe3a947d86d8aa, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test if RandomizedSearchCV is properly defined.
"""

clf_random_cv = make_random_cv()

assert type(clf_random_cv) == RandomizedSearchCV
assert clf_random_cv.get_params()['cv'] == 5
assert clf_random_cv.get_params()['n_iter'] == 20
```

<!-- grade_id: cell-5027e9ebf5206d7a, points: 2, grade: true, solution: false -->
```python
"""
2 point for F1 score higher than .55
Tests are hidden since you don't have access to test targets.
Since there is a random component to Random Search, 
your model will be fit three times and the highest test performance will be used for grading.

Note: There are no tests for Estimator type.
"""

```

<!-- grade_id: cell-9e06df86a051a752, points: 1, grade: true, solution: false -->
```python
"""
1 point for F1 score higher than .65
Tests are hidden since you don't have access to test targets.
Since there is a random component to Random Search, your model will be fit three times and highest test will be use for grading.
"""

```

<!-- grade_id: cell-281369fa7b575435, points: 1, grade: true, solution: false -->
```python
"""
1 point for F1 score higher than .75
Tests are hidden since you don't have access to test targets.
Since there is a random component to Random Search, your model will be fit three times and highest test will be use for grading.
"""

```

Hints:
-----    

- If an algorithm takes a random seed, set it to `random_state=42`.
- Warnings are fine. Errors are not.
- Be careful:
    - Some algorithms have high bias.
    - Some algorithms might perform well on training dataset (low bias), then do poorly on test dataset (high variance). This is called overfitting.
- Review course notes from MSDS 621 and MSDS 699 for ideas.