<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Assignment-#5:-Clustering" data-toc-modified-id="Assignment-#5:-Clustering-1">Assignment #5: Clustering</a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-2">Learning Outcomes</a></span></li><li><span><a href="#Data-Science-Workflow:" data-toc-modified-id="Data-Science-Workflow:-3">Data Science Workflow:</a></span></li><li><span><a href="#1)-Ask:-" data-toc-modified-id="1)-Ask:--4">1) Ask: </a></span></li><li><span><a href="#2)-Acquire" data-toc-modified-id="2)-Acquire-5">2) Acquire</a></span></li><li><span><a href="#3)-Process" data-toc-modified-id="3)-Process-6">3) Process</a></span></li><li><span><a href="#Only-Stolen-Property-incidents" data-toc-modified-id="Only-Stolen-Property-incidents-7">Only <code>Stolen Property</code> incidents</a></span></li><li><span><a href="#4)-Model-" data-toc-modified-id="4)-Model--8">4) Model </a></span></li><li><span><a href="#5)-Deliver" data-toc-modified-id="5)-Deliver-9">5) Deliver</a></span></li><li><span><a href="#Repeat-for-Forgery-And-Counterfeiting-incidents" data-toc-modified-id="Repeat-for-Forgery-And-Counterfeiting-incidents-10">Repeat for <code>Forgery And Counterfeiting</code> incidents</a></span></li><li><span><a href="#Hints-&amp;-Guidelines" data-toc-modified-id="Hints-&amp;-Guidelines-11">Hints &amp; Guidelines</a></span></li></ul></div>

Assignment #5: Clustering
-----

<center><img src="https://img.huffingtonpost.com/asset/5be379c8240000490299a86f.jpeg?ops=scalefit_630_noupscale" width="65%"/></center>

Learning Outcomes
------

__By the end of this session, you should be able to__:


- Fit Gaussian Mixture Model (GMM) to cluster San Francisco Police Department incidents by location.
- Find the number of clusters at that best generalize to unseen data.


Data Science Workflow:
----
1. Ask
2. Acquire
3. Process
4. Model
5. Deliver 

1) Ask: 
----

Can we find meaningful dense areas of Police Department Incidents in San Francisco?

2) Acquire
----

The data comes from https://data.sfgov.org/Public-Safety/Police-Department-Incident-Reports-2018-to-Present/wg3w-h783.

Dataset overview and data dictionary are [here](https://bit.ly/2x7Ta2P).

We are only using the following features:

- Incident Category
- Incident Subcategory
- Latitude
- Longitude

3) Process
-----

```python
reset -fs
```

```python
import pandas as pd
from sklearn.mixture import GaussianMixture
```

```python
path = "../../data/clustering/"

# Load data
data_train = pd.read_csv(path+"data_train.csv")

# Since clustering is unsupervised machine learning, there are no targets.
```

```python
def filter_dataframe(dataframe, incident_category):
    # Select only rows that have the incident_category.
    dataframe = dataframe[data_train['Incident Category'] == incident_category]
    # Select only columns that are lat and long.
    dataframe = dataframe[['Latitude', 'Longitude']]   
    return dataframe
```

Only `Stolen Property` incidents
----

```python
incident_category = "Stolen Property"
X_train_stolen_property = filter_dataframe(dataframe=data_train, incident_category=incident_category)
```

4) Model 
-----

Clustering with Gaussian Mixture Model (GMM)

<!-- grade_id: cell-a1d109c23226e1cb, points: 0, grade: false, solution: true -->
```python
# Find the best number of components for a GaussianMixture as measured by BIC.
# Return a single integer. 
# Do NOT return search code. 

n_components = 50 # TODO: Replace this value with yours.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-91267fcf3b9c2e09, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'n_components' variable. 
This cell should NOT give any errors when it is run.
"""
assert type(n_components) == int
assert n_components > 0
```

```python
# Train model
gmm_stolen_property = GaussianMixture(n_components=n_components, random_state=42).fit(X_train_stolen_property)
# Evaluation metric performance
print(f"{gmm_stolen_property.bic(X_train_stolen_property):,.0f} bic on train dataset.")
```

5) Deliver
-----

How well can your model predict the test data?

<!-- grade_id: cell-1f2c9fcf3c06a34f, points: 1, grade: true, solution: false -->
```python
"""
1 point for bic lower than -2046 on test data.
Tests are hidden since you don't have access to test data.
"""

```

<!-- grade_id: cell-a565bf2f59999160, points: 1, grade: true, solution: false -->
```python
"""
1 point for bic lower than -2273 on test data.
Tests are hidden since you don't have access to test data.
"""

```

<!-- grade_id: cell-46e44db31aad49ad, points: 1, grade: true, solution: false -->
```python
"""
1 point for bic lower than -2417 on test data.
Tests are hidden since you don't have access to test data.
"""

```

<!-- grade_id: cell-1111fc8af80caf0c, points: 1, grade: true, solution: false -->
```python
"""
1 point for bic lower than -2506 on test data.
Tests are hidden since you don't have access to test data.
"""

```

Repeat for `Forgery And Counterfeiting` incidents
-----

```python
incident_category = 'Forgery And Counterfeiting'
X_train_forgery = filter_dataframe(dataframe=data_train, incident_category=incident_category)
```

<!-- grade_id: cell-70a3924a6411eb25, points: 0, grade: false, solution: true -->
```python
# Find the best number of components for a GaussianMixture as measured by BIC.
# Return a single integer. 
# Do NOT return search code. 

n_components = 50 # TODO: Replace this value with yours.

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed64aa1c6484420a, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test code for the 'n_components' variable. 
This cell should NOT give any errors when it is run.
"""
assert type(n_components) == int
assert n_components > 0
```

```python
# Train model
gmm = GaussianMixture(n_components=n_components, random_state=42).fit(X_train_forgery)
# Evaluation metric performance
print(f"{gmm.bic(X_train_forgery):.0f} bic on train dataset.")
```

<!-- grade_id: cell-a360f63a6e60fd20, points: 1, grade: true, solution: false -->
```python
"""
1 point for bic lower than -1905 on test data.
Tests are hidden since you don't have access to test data.
"""

```

<!-- grade_id: cell-bf1248a013dfa711, points: 1, grade: true, solution: false -->
```python
"""
1 point for bic lower than -2093 on test data.
Tests are hidden since you don't have access to test data.
"""

```

<!-- grade_id: cell-e988cac91d89401e, points: 1, grade: true, solution: false -->
```python
"""
1 point for bic lower than -2183 on test data.
Tests are hidden since you don't have access to test data.
"""

```

<!-- grade_id: cell-37a97d3e3cffba5b, points: 1, grade: true, solution: false -->
```python
"""
1 point for bic lower than -2288 on test data.
Tests are hidden since you don't have access to test data.
"""

```

Hints & Guidelines
-----

- Write code to search for number of clusters in another notebook. That notebook is not a deliverable.
- Hints to find best number of clusters:
    - Visual inspection of results.
    - Write code to find minimum.
    - Think about the domain. 
    - Think about how clustering generalizes.