<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#Assignment-#2:-Binary-Classification" data-toc-modified-id="Assignment-#2:-Binary-Classification-0.1">Assignment #2: Binary Classification</a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-0.2">Learning Outcomes</a></span></li></ul></li><li><span><a href="#Data-Science-Workflow:" data-toc-modified-id="Data-Science-Workflow:-1">Data Science Workflow:</a></span><ul class="toc-item"><li><span><a href="#1)-Ask:-" data-toc-modified-id="1)-Ask:--1.1">1) Ask: </a></span></li><li><span><a href="#2)-Acquire" data-toc-modified-id="2)-Acquire-1.2">2) Acquire</a></span></li><li><span><a href="#3)-Process" data-toc-modified-id="3)-Process-1.3">3) Process</a></span></li><li><span><a href="#4)-Model" data-toc-modified-id="4)-Model-1.4">4) Model</a></span></li><li><span><a href="#Process-&amp;-Model-Redux" data-toc-modified-id="Process-&amp;-Model-Redux-1.5">Process &amp; Model Redux</a></span></li><li><span><a href="#5)-Deliver" data-toc-modified-id="5)-Deliver-1.6">5) Deliver</a></span></li><li><span><a href="#Helpful-Hints" data-toc-modified-id="Helpful-Hints-1.7">Helpful Hints</a></span></li></ul></li></ul></div>

Assignment #2: Binary Classification
----
<br>
<br>
<center><img src="https://i.imgur.com/l2nHgCi.png" width="75%"/></center>

Learning Outcomes
------

__By the end of this assignment, you should be able to__:

- Practice applying the Data Science worklflow to a real-world classification problem.
- Fit a binary classification model with scikit-learn.
- Handle modeling unbalanced data.

Data Science Workflow:
======

1. Ask
2. Acquire
3. Process
4. Model
5. Deliver 

1) Ask: 
----

Can a logistic regression model learn to predict extremely unbalanced data (i.e., fraud / not-fraud for credit card transactions)?

2) Acquire
----

The data was sourced from credit cards transactions.

3) Process
-----

```python
reset -fs
```

```python
import numpy as np
import pandas as pd
```

```python
data_dict = {"Variable": "Description",
            "Time": "Time elapsed since first transaction",
            "V1 … V28": "Non-descriptive variables resulting from a PCA dimensionality reduction to protect sensitive data.",
            "Amount": "Transaction Amount",
            "Class": {1:"Fraud",
                      0:"Not Fraud"}}
```

`V1 … V28` are the result of a [PCA](https://en.wikipedia.org/wiki/Principal_component_analysis) transformation of the raw data. 

PCA was done to protect the confidentiality of the transactions.

```python
path = "../../data/binary_classification/"

# Load data
X_train = pd.read_csv(path+"X_train.csv")
y_train = pd.read_csv(path+"y_train.csv")
y_train = y_train.values.ravel()

# Keep track of the number of instances
n_rows_input = X_train.shape[0]
```

<!-- grade_id: cell-357da373ef16d883, points: 0, grade: false, solution: true -->
```python
def process_data(data):
    """Feature engineering & feature selection
    The input parameter and return value have the same name. All processing should happen inplace.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return data
```

<!-- grade_id: cell-c20e0d0ba9fc778b, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the 'process_data' function. 
This cell should NOT give any errors when it is run.
"""

X_train = process_data(X_train)

# Double check the type is still a pd.DataFrame
assert type(X_train) == pd.core.frame.DataFrame
# Double check no rows where drop. Dropping which will break test data performance.
assert X_train.shape[0] == n_rows_input 

```

4) Model
-----

Only use scikit-learn's logistic regression. No other models are allowed.

No automatic hyperparameter search is allowed in final submitted code, thus no model with CV suffix.

```python
from sklearn.linear_model import LogisticRegression
```

<!-- grade_id: cell-f7a79a6df0554d92, points: 0, grade: false, solution: true -->
```python
# Fit Logistic Regression
clf = LogisticRegression() # TODO: Replace with your choice of hyperparameters 
clf.fit(X=X_train, y=y_train) # Train model

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c5eda50206105a66, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test for valid model type.
This cell should NOT give any errors when it is run.
"""

assert "LogisticRegression" in str(type(clf)).split(".")[-1]
assert str(type(clf)).split(".")[-1].count('CV') == 0
```

```python
# Predict on training
y_train_predicted = clf.predict(X_train)
```

```python
# Let's look at accuracy on the training dataset
from sklearn.metrics import accuracy_score

print(f"{accuracy_score(y_train_predicted, y_train):.4f} accuracy on training dataset.")
```

```python
# Let's look at recall on the training dataset
from sklearn.metrics import recall_score

print(f"{recall_score(y_train, y_train_predicted):.4f} recall on training dataset.")
```

```python
# Let's look at the classification report on the training dataset.
from sklearn.metrics import classification_report

print(classification_report(y_train, y_train_predicted))
```

```python
# Let's look at the confusion matrix on the training dataset.
from sklearn.metrics import confusion_matrix

# Just printing raw numbers. Go look at scikit-learn docs for what they mean.
# https://scikit-learn.org/stable/modules/generated/sklearn.metrics.confusion_matrix.html
print(confusion_matrix(y_train, y_train_predicted))
```

```python
# Let's bundle those up for later

def report(y_true, y_pred):
    "A collection of useful metrics for this assigment"
    print(f"{accuracy_score(y_true, y_pred):.4f} accuracy.")
    print(f"{recall_score(y_true, y_pred):.4f} recall.")
    print(classification_report(y_true, y_pred))
    print(confusion_matrix(y_true, y_pred))
```

Process & Model Redux
-------

Often times, you need to repeat the Process & Model steps.

Let's rebalance the data with the [imbalanced-learn package](https://imbalanced-learn.readthedocs.io/en/stable/).

```python
import imblearn
```

```python
# Let's create a validation set 
# The val_train will be rebalanced. The val_test will be left alone as a proxy for the real testdataset.
from sklearn.model_selection import train_test_split

X_val_train, X_val_test, y_val_train, y_val_test = train_test_split(X_train, y_train, test_size=0.30, stratify=y_train)
```

<!-- grade_id: cell-b8cbc8b11b5981c2, points: 0, grade: false, solution: true -->
```python
rebalance_instance = None # TODO: Create your own rebalance_instance

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9bc1994e5f5e460e, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test for valid rebalance type.
This cell should NOT give any errors when it is run.
"""

assert "imblearn" in str(type(rebalance_instance)).split(' ')[1].split(".")[0]
```

```python
X_val_train_smote, y_val_train_smote = rebalance_instance.fit_sample(X_train, y_train)
```

<!-- grade_id: cell-981fa4c12d3311d7, points: 0, grade: false, solution: true -->
```python
# Fit Logistic Regression
clf_sm = LogisticRegression() # TODO: Replace with your choice of hyperparameters 
clf.fit(X=X_val_train_smote, y=y_val_train_smote) # Train model

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a6af818a79fa9508, points: 1, grade: true, solution: false -->
```python
"""
1 point
Test for valid model type.
This cell should NOT give any errors when it is run.
"""

assert "LogisticRegression" in str(type(clf)).split(".")[-1]
assert str(type(clf)).split(".")[-1].count('CV') == 0
```

```python
print("SMOTEd training dataset:")
y_predicted = clf.predict(X_val_train_smote)
report(y_true=y_val_train_smote, y_pred=y_predicted)
```

```python
# See how much we overfit using SMOTE
print("SMOTEd test dataset:")
y_predicted = clf.predict(X_val_test)
report(y_true=y_val_test, y_pred=y_predicted)
```

5) Deliver
-----

How well can your model predict the test data?

<!-- grade_id: cell-2979c840696051f0, points: 2, grade: true, solution: false -->
```python
"""
2 points
Test code for the 'process_data' function. 
This cell should NOT give any errors when it is run.
"""

X_test = pd.read_csv(path+"X_test.csv")
# You don't have access to test targets.

# Keep track of the number of instances
n_rows_input = X_test.shape[0]

X_test = process_data(X_test)

# Double check the type is still a pd.DataFrame
assert type(X_test) == pd.core.frame.DataFrame
# Double check no rows where drop. Dropping which will break test data performance.
assert X_test.shape[0] == n_rows_input 
```

<!-- grade_id: cell-79420c28557d0702, points: 1, grade: true, solution: false -->
```python
"""
1 point for recall score higher than .65
Tests are hidden since you don't have access to test targets.
"""


```

<!-- grade_id: cell-4b64f18cddab90a1, points: 1, grade: true, solution: false -->
```python
"""
1 point for recall score higher than .75
Tests are hidden since you don't have access to test targets.
"""

```

<!-- grade_id: cell-70b1913d033c637c, points: 1, grade: true, solution: false -->
```python
"""
1 point for recall score higher than .85
Tests are hidden since you don't have access to test targets.
"""

```

Helpful Hints
------

- The primary deliverables are `process_data` function and a single logistic regression model, aka `clf`. There should be no other code in the notebook, including visualizations or hyperparameter search code.
- Suggested steps (deliverable / nondeliverable)
    1. Exploratory data analysis (EDA) In particular, understand the class imbalances. (nondeliverable). 
    2. Process data to create features (deliverable)
    3. Search for model, hyperparamter, parameters. I suggest trying out cross-validation. (nondeliverable)
    4. Select single algorithm with single set of hyperparamter (deliverable)
- There are __no__ tricks to doing well on the test dataset. This assignment is about handling class imbalances.
- Use your EDA skills. Create a separate notebook for EDA.
- Think deeply about the data. What potential issues will arise from class imbalances?
- Warnings are fine. You can fix them if you want but they won't effect your grade.
- Errors will negatively effect your grade.
- Academic Integrity reminder - Use only course notes and package (scikit-learn, pandas, …) documentation for this assignment. You should __not__ be googling for help. You should be working independently.